# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print

```julia
println("I'm excited to learn Julia!")
```

Output:
```
I'm excited to learn Julia!
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

Output:
```
String
```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

Output:
```
1
```

```julia
typeof(😺)
```

Output:
```
Int64
```

Note: Julia allows us to write super generic code, and 😺 is an example of this. 

This allows us to write code like

```julia
😀 = 0
😞 = -1
```

Output:
```
-1
```

```julia
😺 + 😞 == 😀
```

Output:
```
true
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

```julia
quotient = 100 / 10
```

Output:
```
10.0
```

```julia
power = 10 ^ 2
```

Output:
```
100
```

```julia
modulus = 101 % 2
```

Output:
```
1
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
#= 
The convert function is convert(T,x) or T(x) that converts x to a value of type T
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia
days = 365 
days_float = float(days)
```

Output:
```
365.0
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
convert(Int64, '1')
parse(Int64, '1')
```

Output:
```
1
```

```julia
Please click on `Validate` on the top, once you are done with the exercises.
```

Output:
```
syntax: extra token "click" after end of expression
```